# NLP and the Web

## Task 0 ~ 0P

__Task 0 is only relevant for the homework.__

##### a) Please enter your group number

<!-- grade_id: cell-cccfbf605a28a18a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

_**Regarding types, documentation, and output:**_

_We tried to make the description of the parameters as clear as possible. However, if you believe that something is missing, please reach out to us in Moodle. We provide type hints for the function parameters and return values of the functions that you have to implement. These are suggestions only, and you may use different types if you prefer._

_Nevertheless, your code must use the provided method stubs and parameters. Furthermore, make sure that your code runs without errors and in a reasonable amount of time, for example by using "Kernel/Restart & Run All" before submitting._

_Please use comments where appropriate to help the tutors understand your code. This is especially important for the more extensive exercises later on. Finally, please pay attention to how you output the results. We highly recommend using `display(df)` for displaying data frames._

_**Please only modify the template in the specified markdown and code cells (e.g. YOUR CODE / ANSWER / IMPORTS HERE). Some cells are left blank on purpose. Please do not modify these cells, because they are used to autograde your submission. If these cells are modified, the automatic grading for your submission will fail and we might deduct points. Please do not modify the cells containing public and private tests. If you want to do your own tests, please use the code cell containing your code solution (YOUR CODE HERE).**_

---

# Tutorial 1

# Theory

### Tokenization

Next, we will describe and show some tokenization ambiguity.
We will present some sentence examples that should be tokenized.

1. Let's meet at 2:30 pm

<!-- grade_id: cell-980e74383a58141e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

2. He's a 6-foot-tall basketball player.

<!-- grade_id: cell-0bb7655db94756b7, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Stemming and Lemmatization

In this task, we would like to dive and gain deeper understanding the difference between stemming and lemmatization. Please give each word after stemming and lemmatization

1. Caring 

<!-- grade_id: cell-94b2a2e48d7bfccd, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

2. Meanness

<!-- grade_id: cell-7ffdba2519e0f6bb, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

3. Better

<!-- grade_id: cell-5335d0d9869fe89a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

# Numpy and Pandas

## Now is the turn for some Numpy insights
Numpy is a python library for mathematical purposes and with good support for vectors and matrices. A library can be loaded into your file using the *import* instruction. We can define an abbreviation for used libraries, so that we do not need to type the full name.

```python
import numpy as np

a = np.arange(15)
print(a)
```

```python
a = a.reshape(5,3)
print(a)
```

### Basics
The main object of numpy is the homogeneous multidimensional array. It is a table of elements of the same type, indexed by non-negative integers. 

Numpy's array class is called ndarray. An array of the type ndarray has several attributes: ndim, shape, size, dtype are the most important ones.

__Question__: What is actually the difference between them? Discuss this with your neighbor.

```python
print(f"dimension in a is: {a.ndim}")
print(f"shape of a is: {a.shape}")
print(f"size of a is: {a.size}")
print(f"type of a: {a.dtype}")
```

### Creation
There are several ways to create a ndarray. To explore them have a look at https://numpy.org/doc/stable/user/quickstart.html. This easiest way is to use np.array:

```python
b = np.array([2,3,5])
print(b)
```

```python
c = np.array([[1,2,3],[4,5,6]])
print(c)
```

It creates a ndarray with the passed Python list and automatically selects the dtype (dont forget the brackets [], this function expects a single Python list as argument). We saw another way to create a numpy array above: *arange()*. 

__Task 1__: Define numpy arrays, containing 10 Zeros, 10 Ones and 10 Fives. (Note: There are more than one solutions)

<!-- grade_id: cell-0ad9718d9f0c359f, points: 0, grade: false, solution: true -->
```python
# Define numpy arrays, containing 10 Zeros, 10 Ones and 10 Fives. (Note: There are more than one solutions)
result = np.array([])

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c876b0f21281b817, points: 0, grade: true, solution: false -->
```python
# Test
assert np.count_nonzero(result == 0) == 10
assert np.count_nonzero(result == 1) == 10
assert np.count_nonzero(result == 5) == 10
```

If an array is too large to be printed, NumPy automatically skips the central part of the array and only prints the corners:

```python
print(np.arange(10000))
```

To disable this behaviour and force NumPy to print the entire array, you can change the printing options using set_printoptions.

```python
import sys
np.set_printoptions(threshold=sys.maxsize)  
```

### Basic Operations
We can also use the ordinary mathematical operations on vectors, i.e. Numpy-arrays. But the vectors must have the correct shape and be of the same type (in some cases casts are possible).

```python
a = np.array([10, 20, 30, 40])
b = np.arange(4)
```

```python
print(a)
print(b)
```

```python
print(a + b)
```

```python
print(a - b)
```

__Question__: What is the difference between the following instructions? Discuss this with your neighbor.

```python
print(a * 2)
print(a * b)
print(a.dot(b))
```

__Task 2__: Define one function, that expects a numpy-matrix and results the sum of all elements in this matrix, the sum of each columns as well as the sum of each row (each as vector). The result for<br>
[[1, 2, 3],<br>
[4, 5, 6],<br>
[7, 8, 9]]<br>
should be 45, [12, 15, 18], [6, 15, 24]. (Hint: *np.sum()*)

<!-- grade_id: cell-ee7b405d1f9a84aa, points: 0, grade: false, solution: true -->
```python
matrix = np.array([[1, 2, 3],
[4, 5, 6],
[7, 8, 9]])
def cal_sum(matrix):
    result0 = 0
    result1 = np.array([])
    result2 = np.array([])
    # YOUR CODE HERE
    raise NotImplementedError()
    return result0, result1, result2

```

<!-- grade_id: cell-26af0368f5cf5d06, points: 0, grade: true, solution: false -->
```python
# Test
result = cal_sum(matrix)
assert result[0] == 45
assert np.all(result[1] == np.array([12, 15, 18]))
assert np.all(result[2] == np.array([6, 15, 24]))
```

### Indexing
After creating an array, we can still assign values to specific positions in the array:

```python
a = np.zeros(10)
a[5] = 1
print(a)
```

For an array with multiple dimensions, the first brackets accesses the corresponding row. To change a value, we must specify which element of that row you want to change:

```python
a = a.reshape(2,5)
print(a)
a[1][4] = 5
print(a)
```

We can easily index the last element of an array with the index -1:

```python
print(a[-1])
```

We can also assign whole columns:

```python
a[:,3] = 1000
print(a)
```

__Task 3__: We want the first row of 'a' to contain '2'.

<!-- grade_id: cell-8119411d42765262, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-92dea5e1ef446e3f, points: 0, grade: true, solution: false -->
```python
# Test
assert np.any(a[0]==2)
```

__Task 4__: Compute the following matrix multiplications: A \* B and B \* A. What is the result of using the ordinary multiplication operator and what do we have to change for using the multiplication operator? (Hint: *array.reshape()*) 

```python
A = a
B = np.ones(10)

A.shape
```

<!-- grade_id: cell-fd99026548746552, points: 0, grade: false, solution: true -->
```python
A = a
B = np.ones(10)

# Output a matrix shape like A
result0 = np.zeros((2,5))
# YOUR CODE HERE
raise NotImplementedError()

result1 = np.zeros(10)
# Output a matrix shape like B
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-69d49b1f617f8cff, points: 0, grade: true, solution: false -->
```python
# Test
assert result0.shape == (2,5)
assert result1.shape == (10,)
```

__Task 5__:: Convert math formulars to numpy，for example sigmoid-function, and print the index which has the biggest value
$$S(x) = \frac{1}{1+e^{(-x)}}$$
![avatar](https://upload.wikimedia.org/wikipedia/commons/thumb/5/53/Sigmoid-function-2.svg/2880px-Sigmoid-function-2.svg.png)

<!-- grade_id: cell-0659aeb889915377, points: 0, grade: false, solution: true -->
```python
np.random.seed(0)
a = np.random.rand(4)
print(a)
# YOUR CODE HERE
raise NotImplementedError()
```

For more informations about numpy have a look at the documentations and tutorials: https://numpy.org/doc/stable/user/quickstart.html

## Pandas

Pandas is a tool you will need frequently in this course. Dataframes, in particular, are used to structure the input and ouput text. Your task in this section is to understand what is going on, as you will need that in the homework exercise.


```python
import pandas as pd
import numpy as np
```

### Creating Dataframes
This first simple way to create and initialize a dataframe is to use a dictionary. Each key in the dictionary is a new column in the dataframe and each column has a different datatype:

For better clarity use the *display()* function to print a dataframe.



```python
df = pd.DataFrame(
    {
        "A": 1.0,
        "B": pd.Timestamp("20130102"),
        "C": pd.Series(1, index=list(range(4)), dtype="float32"),
        "D": np.array([3] * 4, dtype="int32"),
        "E": pd.Categorical(["test", "train", "test", "train"]),
        "F": "foo",
    }
)
display(df)
print(f"the type of df is: \n{df.dtypes}")
```

We can also create an empty dataframe and specify only the column and row names:

```python
df = pd.DataFrame(columns=['Matrikelnummer', "Email", "Address"], index=range(10))
df
```

Now we read a csv file and write it to a new dataframe for example purposes. We will get into more details later.

```python
df = pd.read_csv("yelp_polarity.txt", sep="\t", header=None)
## we can use to_csv methond to save df as csv
df.to_csv('df_test.csv', sep='\t')
display(df)
```

### Viewing data

We can use *head()* and *tail()* to view the top and bottom rows of the frame respectively:

```python
df.head()
```

```python
df.tail()
```

To get the column names and their types we can use the *columns* attribute:

```python
df.columns
```

We also can use this attribut to rename the columns to make them easier to understand:

```python
df.columns= ["Sentence", "Attitude"]
df
```

### Getting/Selection elements of the dataframe

To get the value of a column we can treat the dataframe like an array. The datatype of this column is a Pandas Series. You can find more information about this on their website.

```python
column = df["Sentence"]
print(column) # also df.sentence
print(f"Type of the columns' value: {type(column)}")
```

We can also get 2 columns of a dataframe, but then we get a new dataframe instead of a series.

```python
display(df[["Attitude","Sentence"]]) # also df.loc[:,["Attitude","Sentence"]]
print(type(df[["Attitude","Sentence"]]))
```

To get the row of a dataframe, we can define a range of indices that we want to have. For example, we want to see entries 100-120. The output of this is a new dataframe.

```python
display(df[100:120])  # also df.loc[100:120,:] 
```

Notice that we don't change the original dataframe. We create a new one:

```python
display(df)
```

If we want to have a specific entry by a row and a column, we can use this:

Hint: We can use the column names (Sentence/Attitude) with the *loc()* function or the column indizes (0/1) with the *iloc()* function.

```python
df.iloc[2,0]
```

```python
df.loc[2,"Sentence"]
```

We can also use a boolean expression applied to the dataframe. In this example, we print which entries have an attitude greater than 0.

```python
df["Attitude"] > 0
```

If we want to get all entries that match this boolean expression, we can do it this way:

```python
df[df["Attitude"] > 0]
```

### Methods

# Theory

We can use the *df.apply()* function to apply a function specified in the parameter to every entry in a specified column.

__Task 6:__ use the *df.apply()* function to add 1 to the attitude of all entries. It needs a function as parameter. Use the documentation if you have problems.

<!-- grade_id: cell-4621d58d066e6e36, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-918de04dfd269cbe, points: 0, grade: true, solution: false -->
```python
# Test
assert df["Attitude"][0] == 2 
```

## Text Processing

Now let's go into how to extract information from a text. 

First we will import some packages that we need for our task and here we go!

```python
import pandas as pd
import numpy as np
from sklearn.feature_extraction.text import CountVectorizer
```

Normally models do not process text but a numerical version of it. Therefore we need to vectorize our data. There are several ways of doing it, you might remember that from our previous presentation. Here are a couple of them:

**Bag of words**: 
Let's imagine that we have four documents:

 1. I like to eat pizza every day
 2. I like to eat pasta every day
 3. I do not like pizza
 4. I do my homework every day
 
__Task 7__: We want to list all words, that our documents consists of (i.e. the vocabulary).
You can check the string function *split()* and the method *sorted()*.

```python
corpus = [
    'I like to eat pizza every day',
    'I like to eat pasta every day',
    'I do not like pizza',
    'I do my homework every day'
]
```

<!-- grade_id: cell-d8543527d0efe686, points: 0, grade: false, solution: true -->
```python
def extract_vocab(corpus):
    result = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
```

<!-- grade_id: cell-dae0232b78b1ac7f, points: 0, grade: true, solution: false -->
```python
result = extract_vocab(corpus)
assert result[0] == 'I'
assert result[1] == 'day'
assert result.count(result[0])==1
```

Having the list, we can now write manually our bag of words for each document:

1. [1, 1, 1, 1, 1, 0, 1, 0, 0, 0, 1, 0]
2. [1, 1, 1, 1, 1, 0, 1, 0, 0, 1, 0, 0]
3. [1, 0, 0, 0, 0, 0, 1, 0, 1, 0, 1, 0]
4. [1, 1, 0, 0, 1, 1, 0, 1, 1, 0, 0, 1]

And tada!! This are the vectors that we would feed in our model.

An easier way to do this, is to use *CountVectorizer()*:

```python
corpus = [
    'This is the first document.',
    'This document is the second document.',
    'And this is the third one.',
    'Is this the first document?']
```

```python
count_vectorizer = CountVectorizer()
X = count_vectorizer.fit_transform(corpus)
print(count_vectorizer.get_feature_names_out())
```

```python
print(X.toarray())
```

```python
valid = ["This is the third documents "]
count_vectorizer.transform(valid).toarray()
```

## Reading and Writing Files

There are several ways of processing files, the most common one is using the `open()` function.

```python
my_file = open("yelp_polarity.txt","r")

line_counter = 0

for line in my_file:
    print(line)
    line_counter += 1
    if line_counter >= 5:
        break

my_file.close()
```

**Note**: when handling very big text files, keeping the whole data in the buffer would be counterproductive and inefficient, therefore it is important to learn how to read one line at a time.

```python
with open("yelp_polarity.txt") as my_file:
    line_counter = 0

    for line in my_file:
        print(line)
        line_counter += 1
        if line_counter >= 5:
            break
```

As I said, if you have files in CVS format, you can read them with Pandas, which is a useful way to analyze your data.

```python
dataframe = pd.read_csv("yelp_polarity.txt", sep="\t", header=None)
```

```python
display(dataframe)
```

__Task 8__: Now it's your turn to explore data using Pandas. How many classes do we have in this file and how many samples per class? Remember that the question mark (?) can be always very useful when you don't know the functions. Check the function `groupby` and `describe`.

<!-- grade_id: cell-73c79b031ef58341, points: 0, grade: false, solution: true -->
```python
# Check the function groupby and describe
# YOUR CODE HERE
raise NotImplementedError()
```