# NLP and the Web

## Task 0 ~ 0P

##### a) Please enter your group number

<!-- grade_id: cell-cccfbf605a28a18a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

_**Regarding types, documentation, and output:**_

_We tried to make the description of the parameters as clear as possible. However, if you believe that something is missing, please reach out to us in Moodle. We provide type hints for the function parameters and return values of the functions that you have to implement._

_Nevertheless, your code must use the provided method stubs and parameters. Furthermore, make sure that your code runs without errors and in a reasonable amount of time, for example by using "Kernel/Restart & Run All" before submitting._

_Please use comments where appropriate to help the tutors understand your code. This is especially important for the more extensive exercises later on. Finally, please pay attention to how you output the results. We highly recommend using `display(df)` for displaying data frames._

_**Please only modify the template in the specified markdown and code cells (e.g. YOUR CODE / ANSWER / IMPORTS HERE). If you add any extra cells, they wont be taken into account while grading!  Some cells are left blank on purpose. Please do not modify these cells, because they are used to autograde your submission. If these cells are modified, the automatic grading for your submission will fail and we might deduct points. Please do not modify the cells containing public and private tests. If you want to do your own tests, please use the code cell containing your code solution (YOUR CODE HERE).**_

---

#  Home Exercise 5

# Imports

```python
import numpy as np
from typing import Tuple, List
import numpy.typing as npt
import matplotlib.pyplot as plt
```

## Task 1: Neural Network from scratch ~ 10P

**a) What is a Perceptron and what are the tasks of the 5 different components? (1P)**

<!-- grade_id: cell-b61af27cdc9fc324, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**b) Which activation functions do you know besides the ReLU (name at least 3) and explain the inputs / outputs of one of those? (1.5P)**

<!-- grade_id: cell-0afb53fc3ad6f345, points: 1.5, grade: true, solution: true -->
YOUR ANSWER HERE

**c) Implement the ReLU activation function and its derivative. (2P)**

*Note:* The ReLU (Rectified Linear Unit) function is a piecewise linear function which is very similar to the linear activation function that you know from lecture 2 page 30. The difference between these two is that ReLU cuts of the negative value and sets them to zero. $$\text{ReLU}(x) = \begin{cases} x, & \text{if } x > 0 \\ 0, & \text{if } x \le 0 \end{cases} $$

<!-- grade_id: cell-59f29cac5e1ab452, points: 0, grade: true, solution: true -->
```python
def relu(z: npt.NDArray[np.float64]) -> npt.NDArray[np.float64]:
    '''
    Input: 
        z: floating point vector to which you want to apply the ReLU activation function

    Output:
        Output a floating point vector that has the same shape as z but with the activated values
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return az
```

```python
relu(np.array([1,-1]))
```

<!-- grade_id: cell-7cc5e0bafd23eda9, points: 1, grade: true, solution: false -->
```python
# These are some test cases for you to check your implementation
assert((relu(np.array([1,-1])) == np.array([1,0])).all())
assert((relu(np.array([-112,234])) == np.array([0,234])).all())
```

<!-- grade_id: cell-25722a6942cadbc3, points: 0, grade: true, solution: true -->
```python
def grad_relu(z: npt.NDArray[np.float64]) -> npt.NDArray[np.float64]:
    '''
    Input: 
        z: Vector to which you want to apply the gradient of the ReLU activation function

    Output:
        Output a vector that has the same shape as z and is the derivative of z
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return d_az
```

<!-- grade_id: cell-a7a9a25aa59363f7, points: 1, grade: true, solution: false -->
```python
# These are some test cases for you to check your implementation
assert((grad_relu(np.array([1,-1])) == np.array([1,0])).all())
assert((grad_relu(np.array([-112,234])) == np.array([0,1])).all())
```

**d) Implement the forward function and the cost function and its gradient (2.5P).**

Implement the forward function for a neural network with one hidden layer.

*Hint:* You only need to use the activation function in the hidden layer.

<!-- grade_id: cell-ddd55abc575278ce, points: 0, grade: true, solution: true -->
```python
def forward(X: npt.NDArray[np.float64], 
            Wh: npt.NDArray[np.float64], 
            Bh: npt.NDArray[np.float64], 
            Wo: npt.NDArray[np.float64], 
            Bo: npt.NDArray[np.float64]) -> Tuple[npt.NDArray[np.float64], 
                                                  npt.NDArray[np.float64], 
                                                  npt.NDArray[np.float64]]:
    ''' 
    This is the forward function for making predictions. 
    In order to train the network later you also need to return the activation and pre-activation of the hidden layer.
    
    Input:
        X: This is equivalent to the x in f(x). So to this value we want a prediction, Dimension: Scalar
        Wh: weight matrix or vector for the hidden layer, Dimension: inputLayerSize x hiddenLayerSize
        Bh: bias for the hidden layer, Dimension: 1 x hiddenLayerSize
        Wo: weight matrix for the output layer, Dimension: hiddenLayerSize x ouputLayerSize
        Bo: bias for the output layer, Dimension: 1 x ouputLayerSize

    Output:
        1. Prediction, Dimension: Scalar
        2. Activation of the hidden layer, Dimension: 1 x hiddenLayerSize
        3. Pre-activation of the hidden layer, Dimension: 1 x hiddenLayerSize
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return Zo, Ah, Zh
    
```

<!-- grade_id: cell-61597567b468dbbf, points: 1.5, grade: true, solution: false -->
```python
# Test - Don't modify
```

The cost function (also known as *loss*) is given as follows: $$\begin{align*} X &:= \text{Input for forward function}\\ y &:= \text{Correct label for } X \\ \text{forward}(X) &:= \text{Forward function that predicts a label for input } X\\ \text{The cost function:}\\ \text{cost}(X,y) &= \frac{\sum (\text{forward}(X)-y)^2}{2} \end{align*} $$

This is a slight variation of the one you already learned in the lecture but will do just fine.

<!-- grade_id: cell-9cd1af9bdd6992c7, points: 0, grade: true, solution: true -->
```python
def cost(X: npt.NDArray[np.float64], y: np.float64,
            Wh: npt.NDArray[np.float64], 
            Bh: npt.NDArray[np.float64], 
            Wo: npt.NDArray[np.float64], 
            Bo: npt.NDArray[np.float64]) -> np.float64:
    ''' 
    Quadratic loss function.
    
    Dimensions are similar to forward function.

    Input:
        X: Input for the value to predict
        y: real value for the prediction, Dimension: Scalar
        Wh: weight matrix or vector for the hidden layer
        Bh: bias for the hidden layer
        Wo: weight matrix for the output layer
        Bo: bias for the output layer

    Output:
        cost, Dimension: Scalar
    '''
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return cost
```

<!-- grade_id: cell-4253b8f8ec4d529a, points: 0.5, grade: true, solution: false -->
```python
# Test - Don't modify

```

<!-- grade_id: cell-634c90ea8da4ecc0, points: 0, grade: true, solution: true -->
```python
def cost_grad(output: np.float64, y: np.float64) -> np.float64:
    ''' 
    Gradient of Quadratic loss function

    Input:
        output: predicted value
        y: real value for the output

    Output:
        gradient of cost
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    return grad
```

<!-- grade_id: cell-6fa874405335c61d, points: 0.5, grade: true, solution: false -->
```python
# Test - Don't modify

```

Make yourself a clear image of what gradients are required and how they are calculated. We implemented the `backprop` function already for you.

```python
def backprop(X: npt.NDArray[np.float64],
             output: npt.NDArray[np.float64],
             y: npt.NDArray[np.float64],
             Ah: npt.NDArray[np.float64],
             Zh: npt.NDArray[np.float64],
             Wo: npt.NDArray[np.float64]) -> Tuple[npt.NDArray[np.float64], 
                                                   npt.NDArray[np.float64], 
                                                   npt.NDArray[np.float64], 
                                                   npt.NDArray[np.float64]]:

    '''
    The backprop function should output the gradients to all the necessary vectors. 
     
    Input:
        X: This is equivalent to the x in f(x). So to this value we want a prediction
        output: Predicted value
        y: true value to X
        Ah: activation of hidden layer
        Zh: pre-activation of hidden layer
        Wo: weights of output layer

    Output:
        1. Gradients of weights of hidden layer
        2. Gradients of weights of output layer 
        3. Gradients of bias of hidden layer
        4. Gradients of bias of output layer
    '''
    err_output: np.float64 = cost_grad(output, y) #last layer is linear, no gradient needed
    err_hidden: npt.NDArray[np.float64] = np.dot(err_output, Wo.T) * grad_relu(Zh)

    # Weight Gradients
    dCdWh: npt.NDArray[np.float64] = np.dot(X, err_hidden)
    dCdWo: npt.NDArray[np.float64] = np.dot(Ah.T, err_output)
    
    # Bias Gradients
    dCdBh: npt.NDArray[np.float64] = np.sum(err_hidden, axis=0, keepdims=True)
    dCdBo: npt.NDArray[np.float64] = np.sum(err_output, axis=0, keepdims=True)

    return dCdWh, dCdWo, dCdBh, dCdBo
```

**e) Implement the training loop and train a simple function in the from of $f(x) = ax² + bx + c$ for example $f(x) = (x-5)^2$. But feel free to experiment and maybe try higher order functions too.**

For that you have to implement *gradient descent* to update the weights and performs one epoch of the training cycle and after that the training loop. (3P)

<!-- grade_id: cell-edb04359f5df9ef7, points: 0, grade: true, solution: true -->
```python
def gradient_descent(x: List[float], 
                     y: List[float],
                     lr: float, 
                     Wh: npt.NDArray[np.float64], 
                     Wo: npt.NDArray[np.float64], 
                     Bh: npt.NDArray[np.float64], 
                     Bo: npt.NDArray[np.float64]) -> Tuple[float, 
                                                           npt.NDArray[np.float64], 
                                                           npt.NDArray[np.float64], 
                                                           npt.NDArray[np.float64], 
                                                           npt.NDArray[np.float64]]:
    ''' 
    Gradient descent gets a list of inputs and true outputs and iterates trough 
    them and performs a weight and bias update each step. 
    
    Dimensions of the varables are similar to task 1 d)

    Input:
        x: list of inputs in the same ordering as y
        y: list of true outputs with the same ordering as x
        lr: learning rate of gradient descent
        Wh: weight matrix or vector for the hidden layer
        Bh: bias for the hidden layer
        Wo: weight matrix for the output layer
        Bo: bias for the output layer
    
    Output:
        1. average of cost function of input data (TODO: check formulation)
        2. updated weights of hidden layer
        3. updated weights of output layer
        4. updated bias of hidden layer
        5. updated bias of output layer
    '''
    costs: float = 0.0
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return costs, Wh, Wo, Bh, Bo
```

```python
### Tests to check your implementation
inputLayerSize_t = 1
hiddenLayerSize_t = 5 # This is the value you can play around with
outputLayerSize_t = 1 

# Initialize Weights
Wh_t = np.ones((inputLayerSize_t, hiddenLayerSize_t)) 
Wo_t = np.ones((hiddenLayerSize_t, outputLayerSize_t))

# Initialize Biases
# Small positive value to avoid Relu units dying too quickly
Bh_t = np.full((1, hiddenLayerSize_t), 0.2)
Bo_t = np.full((1, outputLayerSize_t), 0.2)

# gradient_descent([1,2],[3,4], 0.001, Wh_t, Wo_t, Bh_t, Bo_t)

# Forward inputs and outputs
## Test 1
x_1 = [1,2,3]
y_1 = [3,2,1]
Wh_1 = Wh_t
Wo_1 = Wo_t
Bh_1 = Bh_t
Bo_1 = Bo_t

tmp_1 = gradient_descent(x_1,y_1, 0.001, Wh_1, Wo_1, Bh_1, Bo_1)

# for i in range(5):
#     print(np.round(tmp_1[i], 2))

assert(np.round(tmp_1[0], 2) == np.round(42.57, 2))
assert((np.round(tmp_1[1], 2) == np.round(np.array([[0.94, 0.94, 0.94, 0.94, 0.94]]), 2)).all())
assert((np.round(tmp_1[2], 2) == np.round(np.array([[0.93], [0.93], [0.93], [0.93], [0.93]]), 2)).all())
assert((np.round(tmp_1[3], 2) == np.round(np.array([[0.17, 0.17, 0.17, 0.17, 0.17]]), 2)).all())
assert((np.round(tmp_1[4], 2) == np.round(np.array([[0.17]]), 2)).all())
```

<!-- grade_id: cell-4a7162dbc77a1cd2, points: 2, grade: true, solution: false -->
```python
# Test - Don't modify

```

The following is the initialization of the weights and biases. So you get consistent results. Feel free to experiment with the number of neurons in the hidden layer. 

```python
def init_network() -> Tuple[npt.NDArray[np.float64], npt.NDArray[np.float64], npt.NDArray[np.float64], npt.NDArray[np.float64]]:
    inputLayerSize: int = 1
    hiddenLayerSize: int = 10 # This is the value you can play around with
    outputLayerSize: int = 1 

    # Initialize Weights
    np.random.seed(0)
    Wh: npt.NDArray = np.random.randn(inputLayerSize, hiddenLayerSize) /np.sqrt(inputLayerSize)
    Wo: npt.NDArray = np.random.randn(hiddenLayerSize, outputLayerSize) /np.sqrt(hiddenLayerSize)

    # Initialize Biases
    # Small positive value to avoid Relu units dying too quickly
    Bh: npt.NDArray = np.full((1, hiddenLayerSize), 0.2)
    Bo: npt.NDArray = np.full((1, outputLayerSize), 0.2)

    return Wh, Wo, Bh, Bo
```

Next is the generation of the training data. Here you can also change up the function which the network should learn or the interval in which the numbers are generated. 

*Hint:* Start with an easy function an later aim for something more difficult.

```python
def f_x(x: float) -> float:
    '''
    Feel free to play around!
    '''
    # return 0.1 * ((x-5)**3 - 10 * x + 120) # Example of a more complex function
    return (x-5)**2

x = np.linspace(0, 10, 1001)
y = np.array([ f_x(i)  for i in x])

plt.plot(x,y)
```

Now implement the training loop and train your network for 1000 epochs with a learning rate of 0.001. Shuffle the training data before each epoch.

*Note:* The cost should go down if that is not the case check your code above. Are the derivatives correct?

**At the end the cost for $f(x) = (x-5)^2$ has to be below 0.05!**

<!-- grade_id: cell-5dc46853dbb96276, points: 0, grade: true, solution: true -->
```python
Wh, Wo, Bh, Bo = init_network()
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e277af25361b7e06, points: 1, grade: true, solution: false -->
```python
# Test - Don't modify

```

Normally you also need test data in order to validate the accuracy of you network. But for now we only plot the predictions and the real values in the same plot so we can compare it.

```python
x_ax = np.linspace(0, 10, 1001)
y = np.array([ f_x(i)  for i in x_ax])
preds = np.array([forward(i, Wh, Bh, Wo, Bo)[0][0,0] for i in x_ax])
plt.plot(x_ax, preds)
plt.plot(x_ax,y)
```

## Task 2 - Byte Pair Encoding ~ 10P

The goal is to calculate byte pair encodings for a given dataset. For a recap of the concept we refer to lecture 9, slides 66.ff.

But first we need to install some additional libraries in order to get it working properly:

```python
!pip install datasets==2.16.1
!pip install tokenizers==0.15.0
```

```python
from tokenizers.pre_tokenizers import Whitespace
from datasets import load_dataset
from datasets.arrow_dataset import Dataset
from collections import defaultdict
```

**a)** First we want to make sure we create a suffieciently large corpus. As dataset you import the dataset "wikitext-2-raw-v1" from huggingface. Make sure to use the 500 first entries of the training dataset and subsequently filter out empty lines. 

Save the dataset in the variable `dataset` and the corpus in the variable `corpus`. ***(1P)***

<!-- grade_id: cell-244778a82446be41, points: 0, grade: true, solution: true -->
```python
# Load WikiText-2 dataset "wikitext-2-raw-v1"

dataset : Dataset = None
corpus: list = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e0ff412ce441cbd7, points: 0.25, grade: true, solution: false -->
```python
# Test - Don't modify

```

<!-- grade_id: cell-0e9998b74bc12e96, points: 0.25, grade: true, solution: false -->
```python
# Test - Don't modify

```

<!-- grade_id: cell-5fefde03129fd7c1, points: 0.25, grade: true, solution: false -->
```python
# Test - Don't modify

```

<!-- grade_id: cell-e08087f00de78bdd, points: 0.25, grade: true, solution: false -->
```python
# Test - Don't modify

```

Note that the corpus should have a format like:

```python
corpus = [
    "This is the first line",
    "This is the second line",
    ..
]
``` 

For debugging during coding you can also consider working with a smaller corpus, e.g. only 5 lines.

**b)** Here, you loop through the corpus and count the word frequencies and save the output in the variable `word_freqs`. 

The `Whitespace` pre_tokenizer can be used for preprocessing a string to strip away empty spaces as well as to separate out special characters such as punctuations. 

*Hint:* You may want to have a look at the function: `pre_tokenize_str` of the `Whitespace` pre_tokenizer. This might help you. ***(2P)***

<!-- grade_id: cell-8ce3e5e867205cc4, points: 0, grade: true, solution: true -->
```python
# Initialize a Whitespace pre_tokenizer from the tokenizers library
# Helps in separating text about white space and punctuations
tokenizer = Whitespace()

# Initialize a defaultdict to store word frequencies
word_freqs = defaultdict(int)
# this dictionary gets entries for word_freqs[word]

# This will be needed later
vocab_size:int = 200

def count_word_freqs(corpus:list) -> dict:
    # We loop through to corups to calculate word frequencies
    # YOUR CODE HERE
    raise NotImplementedError()
    return word_freqs

word_freqs = count_word_freqs(corpus)
print(word_freqs)
```

<!-- grade_id: cell-a261f5385cdff84d, points: 2, grade: true, solution: false -->
```python
# Test - Don't modify
assert len(word_freqs.keys()) == 6911
assert sum(word_freqs.values()) == 40051

```

The following block creates the alphabet and initial vocabulary

```python
def initialize(word_freqs:dict) -> Tuple[list, list, dict, dict]:
    alphabet:list = []
    vocab:list = [""]
    splits:dict = {}
    merges:dict = {}
    
    for word in word_freqs.keys():
        for letter in word:
            if letter not in alphabet:
                alphabet.append(letter)
    alphabet.sort()

    vocab = alphabet.copy()
    splits = {word: [c for c in word] for word in word_freqs.keys()}

    return alphabet, vocab, splits, {}

alphabet, vocab, splits, merges = initialize(word_freqs)
```

**c)** Define a function to compute the frequency of each pair of characters ***(2P)***

<!-- grade_id: cell-5e59098a15cbf372, points: 0, grade: true, solution: true -->
```python
def calc_pair_freqs(splits:dict) -> dict:
    pair_freqs = defaultdict(int)
    # to compute the frequency of each pair of characters
    # YOUR CODE HERE
    raise NotImplementedError()
    return pair_freqs

# Initialize pair_freqs by calling the calc_pair_freqs function
pair_freqs = calc_pair_freqs(splits)
print(pair_freqs)
```

<!-- grade_id: cell-a7b55d57116b9963, points: 2, grade: true, solution: false -->
```python
# Test - Don't modify
assert len(pair_freqs.keys()) == 969
assert sum(pair_freqs.values()) == 130300

```

Here we find the pair of characters that appears most frequently together

```python
for i, key in enumerate(pair_freqs.keys()):
    print(f"{key}: {pair_freqs[key]}")
    if i >= 5:
        break

best_pair = ""
max_freq = None

for pair, freq in pair_freqs.items():
    if max_freq is None or max_freq < freq:
        best_pair = pair
        max_freq = freq

print(best_pair, max_freq)
```

**d)** Define a function to merge the most frequent pair of characters ***(2P)***

<!-- grade_id: cell-f31bb69bdadebf4f, points: 0, grade: true, solution: true -->
```python
def combine_pair(a:str, b:str, splits:dict, word_freqs:dict) -> dict:
    for word in word_freqs:
        split = splits[word]
        if len(split) == 1:
            continue

        # YOUR CODE HERE
        raise NotImplementedError()
    return splits
```

<!-- grade_id: cell-6d34502f30668a98, points: 2, grade: true, solution: false -->
```python
# Test - Don't modify
# Example and sanity check
splits_example = combine_pair("y", "r", splits, word_freqs)
assert splits_example["Valkyria"] == ['V', 'a', 'l', 'k', 'yr', 'i', 'a']

```

**e)** Create the vocabulary by computing the pair frequencies and merging them accordingly, until you have a vocabulary size of 200 like it is done in the lecture notes of lecture 9. ***(2P)***

<!-- grade_id: cell-c0d21a5824af85c6, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3c432d2dedb4097a, points: 2, grade: true, solution: false -->
```python
# Test - Don't modify
assert len(vocab) == 200

```

We provide the tokenize function to you

```python
def text_splitting(text:str, merges:dict, splits:dict) -> list:
    tokens = tokenizer.pre_tokenize_str(text)
    tokens_text = [word for word, _ in tokens]
    splits = [[l for l in word] for word in tokens_text]
    for pair, merge in merges.items():
        for ind, split in enumerate(splits):
            i = 0
            while i < len(split) - 1:
                if split[i] == pair[0] and split[i + 1] == pair[1]:
                    split = split[:i] + [merge] + split[i + 2 :]
                else:
                    i += 1
            splits[ind] = split

    return sum(splits, [])

alphabet, vocab, splits, merges = create_vocab(word_freqs, vocab_size)
text_splitting("Darmstadt holds the official title City of Science (German: Wissenschaftsstadt) as it is a major centre of scientific institutions, universities, and high-technology companies.", merges, splits)
```

**f)** Now, repeat the byte pair encoding for vocab sizes of 1000 and 5000. What do you notice? ***(1P)***

```python
alphabet, vocab, splits, merges = create_vocab(word_freqs, 1000)
text_splitting("Darmstadt holds the official title City of Science (German: Wissenschaftsstadt) as it is a major centre of scientific institutions, universities, and high-technology companies.", merges, splits)
```

```python
alphabet, vocab, splits, merges = create_vocab(word_freqs, 5000)
text_splitting("Darmstadt holds the official title City of Science (German: Wissenschaftsstadt) as it is a major centre of scientific institutions, universities, and high-technology companies.", merges, splits)
```

<!-- grade_id: cell-659e803b4ca7bd4d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Submission:**

Please upload your submission to Moodle before the next exercise session <font color="red">(Jan 24, 23:59pm)</font>!

Submission format: `homework 5.zip`

Your submission should contain your filled out Jupyter notebook (naming schema: `homework 5.ipynb`) and any auxiliar files that are necessary to run your code (e.g., the datasets provided by us).

Each submission must be handed in only once per group.